## FHNW Machine Learning - FS2024

```python
# DO NOT EDIT THIS CELL
# ADD IMOPRTS IN ASSIGNMENT CELLS

# Notebook configs
%load_ext autoreload
%autoreload 2
%matplotlib inline

from IPython.core.interactiveshell import InteractiveShell
from IPython.display import HTML

InteractiveShell.ast_node_interactivity = "all"
```

# Bewertete Übung 02

**Ausgabe:** Montag, 15. April 2024

**Abgabe:** Montag, 6. Mai 2024, 23:59 Uhr


## Vorgaben zu Umsetzung und Abgabe

- Die Algorithmen müssen auf der Basis von Array Operationen selber implementiert werden.
- Der Code muss lauffähig sein bei Ausführung im Docker-Container zum Repo oder auf JHub. 
- Es darf kein Code ausgelagert werden, i.e. sämtlicher Code muss sich im Notebook befinden.
- Sämtliche Plots sind komplett beschriftet (Achsen, Labels, Überschrift, Colorbar, ..), so dass der Plot ohne den Code zu konsultieren, verstanden werden kann.
- Als **Abgabe** zählt der letzte Commit vor Abgabetermin in in Ihrem Fork des Repos.  

- **Bitte löschen, kopieren, duplizieren, splitten und verschieben Sie die vorhandenen Zellen nicht**. Dies führt zu Problemen bei der Korrektur. Sie dürfen aber beliebig viele weitere Zellen hinzufügen (nur via **insert new cell**).
- Laufzeit vom Notebook: Das Notebook sollte in weniger als 30 Minuten ausgeführt werden können.

Für die Erarbeitung der Lösung darf unter Studierenden zusammengearbeitet werden. Die Zusammenarbeit ist dabei aber auf konzeptionelle und algorithmische Fragen und Verständnisaspekte beschränkt.  

**Es darf kein Code oder Text von anderen oder vom Internet kopiert werden.**


### Module

Neben den Python-Basismodulen dürfen Sie die folgenden Module immer benutzen: `numpy`, `pandas`, `matplotlib`, `seaborn`,  `tqdm` (für Progress-Bars).

Sie dürfen auch generell [sklearn.preprocessing](https://scikit-learn.org/stable/modules/classes.html#module-sklearn.preprocessing), [sklearn.model_selection](https://scikit-learn.org/stable/model_selection.html), [sklearn.pipeline](https://scikit-learn.org/stable/modules/classes.html#module-sklearn.pipeline) und [sklearn.compose](https://scikit-learn.org/stable/modules/classes.html#module-sklearn.compose) benutzen.

Zusätzliche Module dürfen Sie nur verwenden wenn ausdrücklich erwähnt oder bereits vorgegeben in der Code-Cell.

## Bewertung

Bewertet wird:

- Vollständigkeit (Code, Text)
- Korrektheit (Code, Text)
- Implementation (z.B. Vektorisierung der Operationen, Scikit-Learn API, Visualisierungen, Lesbarkeit Code/Output)

## Einleitung

Mika arbeitet für eine Hotelkette welche mit schlecht ausgelasteten Hotels zu kämpfen hat. In einer Datenanalyse hat man festgestellt, dass ein wichtiger Grund für die schlechte Auslastung, stornierte Buchungen sind. Diese Buchungen werden zum Teil auch kurzfristig storniert und können dann oft nicht mehr mit neuen Gästen kompensiert werden. Das Management des Hotels möchte nun herausfinden, ob es möglich ist vorherzusagen, welche Buchungen storniert werden. Mika hat nun die Aufgabe bekommen ein Modell zu erstellen, um für jede Buchung eine Vorhersage zu erstellen. Mika hat von zwei Hotels Datensätze gesammelt und diese aufbereitet. Mika hat von Alex erfahren, dass Sie ein Experte in Sachen Machine Learning sind und bittet Sie um Hilfe.

Es ist das Ziel aus Buchungsdaten (siehe Tabelle unten) vorherzusagen, ob diese storniert wird `IstStorniert`.


| Spalt                  | Beschreibung |
|---------------------------------|----------------|
| DurchschnittlicherTagespreis    | Durchschnittlicher Preis pro Nacht |
| Erwachsene                      | Anzahl der Erwachsenen |
| Ankunftsmonat                   | Monat der Ankunft |
| Babys                           | Anzahl der Babys |
| Kinder                          | Anzahl der Kinder |
| Anzahlungstyp                   | Typ der Anzahlung |
| IstStorniert                    | Ob die Buchung storniert wurde (0=Nein, 1=Ja) |
| IstWiederholungsgast            | Ob der Gast ein Wiederholungsgast ist |
| Vorlaufzeit                     | Tage zwischen Buchung und Ankunft |
| Marktsegment                    | Marktsegment der Buchung |
| Mahlzeittyp                     | Gebuchter Mahlzeitentyp |
| VorherigeNichtStornierteBuchungen | Anzahl vorheriger nicht stornierter Buchungen |
| VorherigeStornierungen          | Anzahl vorheriger Stornierungen |
| BenötigteParkplätze             | Anzahl benötigter Parkplätze |
| ReservierterZimmertyp           | Typ des reservierten Zimmers |
| AufenthalteAnWochenendnächten   | Anzahl der Übernachtungen am Wochenende |
| AufenthalteAnWochennächten      | Anzahl der Übernachtungen unter der Woche |
| GesamtzahlDerSonderwünsche      | Anzahl der Sonderwünsche |
| Hotel | ID vom Hotel für welches die Buchung gemacht worden ist



## Aufgabe 1 (5 Punkte)

In dieser Aufgabe ist es das Ziel den Datensatz `./hotel_train.csv` einzulesen, kennenzulernen und erste Erkenntnisse zu diskutieren.


Charakterisieren Sie den Datensatz indem Sie folgende Aufgaben erledigten, bzw. Fragen beantworten. Nur die Beantwortung der folgenden Fragen wird bewertet:

1. Wie beeinflusst die Anzahl und die Kategorien der Personen, die eine Buchung umfasst, die Zielvariable `IstStorniert`?
2. Ist der Preis pro Nacht Abhängig von den Anzahl Personen?
3. Schauen Sie sich die Wiederholungsgäste an: Wieviele gibt es und wie häufig stornieren diese im Vergleich zu den anderen Gästen?
4. Von wievielen Hotels stammen die Buchungen? Was könnte das für einen Einfluss haben auf die Modellierung?
5. Unser Ziel ist es den Buchungsstatus `IstStorniert` vorherzusagen unter Verwendung der anderen Variablen.  Was lässt sich über dieses Vorhaben sagen? Diskutieren Sie Ihre Einsichten.

Beantworten Sie alle Fragen in der folgenden Text-Zelle. Referenzieren Sie die Fragen anhand der Fragennummer. Sie können auf Grafiken oder Tabellen referenzieren, die Sie unten in der Code-Zelle erstellen können.


<!-- grade_id: answer1a, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

Erstellen Sie bei Bedarf Grafiken / Outputs um Fragen zu beantworten in der folgenden Zelle.

<!-- grade_id: answer1b, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Aufgabe 2 (23 Punkte)

Implementieren Sie einen _Random Forest Classifier_ basierend auf binären _Decision Trees_. Folgende Eigenschaften sollen vorhanden sein:

- Man soll ein binäres Klassifikationsproblem modellieren können ($y \in [0, 1]$).

- Konfigurierbarkeit bezüglich vorgegebenen Parametern wie Abbruchkriterien (Regularisierung).

- Korrektes bzw. sinnvolles Handling von verschiedenen Aspekten wie z.B. das Verhindern von leeren _Leaf Nodes_.

- Modulare Implementation: Erstellen / Ergänzen Sie die bereitgestellten Klassen. Sie können auch neue Methoden oder Argumente hinzufügen.

- Reproduzierbarkeit: Setzen Sie _random seeds_ wo nötig [siehe numpy](https://numpy.org/doc/stable/reference/random/index.html#random-quick-start).
- score() soll die Accuracy berechnen.

- Konsultieren Sie die Unterrichtsunterlagen, insbesondere für die Berechnung / das Finden der optimalen Splits.

- Die Implementation benötigt viele Loops und lässt sich nicht gut vektorisieren.

- Implementieren Sie das Modell gemäss sklearn-API.

Im folgenden einige Hinweise.


Die Vorgaben der Implementation sind sehr nahe an dem was in den Kursunterlagen gezeigt wird. Konsultieren Sie diese bei Bedarf.

### Decision Tree

Der Random Forest Classifier besteht aus einem _Ensemble_ von _Decision Tree Classifiern_. Die wichtigsten Funktionalitäten werden also in erstern Linie in der Klasse `BinaryDecisionTreeClassifier` implementiert sein. Die `BinaryRandomForestClassifier` Klasse managt das Training und das Ausführen von _Decision Trees_. 

Sie sollen einen `BinaryDecisionTreeClassifier` implementieren, der ausschliesslich binäre Splits auf nummerischen Variablen durchführt (das beinhaltet auch kategorielle Variablen die entsprechend _pre-processed_ worden sind). Ein solcher Decision Tree ist folgendermassen aufgebaut:

![Binärer Decision Tree](decision_tree.png)


### Random Forest

Ein Random Forest ist eine Kollektion von Decision Trees. Wobe jeder Decision Tree auf einem Trainingsdatensatz trainiert wird welcher mit _Resampling with Replacement_ erstellt worden ist.

Ein Random Forest funktioniert also im wesentliche so:

```
def fit(self, X, y):
    my_trees = list()
    for i in num_trees:
        X_sample, y_sample = sample_with_replacement(X, y)
        dt = DecisionTreeClassifier(x_sample, y_sample)
        my_trees.append(df.fit())
```

Zusätzlich werden bei jedem Node von einem Decision Tree nur ein zufälliges Sample von $l \le k$ Features berücksichtigt (statt alle $k$ Features zu berücksichtigen).

```
def fit(self, X, y):
    ...

    node_features = sample(l, all_features)

    best_feature = find_best_feature(node_features)
    
    ...
```

Um Predictions zu generieren werden alle Decision Trees konsultiert. Dabei wird die Klasse mit den meisten Stimmen (_Majority Vote_) als Vorhersage verwendet. Diese wird dann dem Datenpunkt zugeordnet.


### Building a Decision Tree

Ein Decision Tree wird rekursiv erstellt. Man erstellt als erstes den Root-Node und splittet diesen dann. Dabei wird die gleiche Logik rekursiv auf alle Child-Nodes angewendet, bis ein Stop-Kriterium erreicht worden ist. 

Man sucht für jeden Node den optimalen Split $S_{j,t}$. Dieser ist definiert durch welches Feature $j$ bei welchem Threshold $t$ der Datenatz im entsprechenden Node (binär) gesplittet wird.

#### Impurity

Man möchte, dass die Nodes von einem Decision Tree möglichst uniform sind (tiefe _Impurity_ haben), d.h. möglichst aus Datenpunkten von einer Klasse bestehen. Man minimiert dazu eine Impurity Metrik $Q$ . Wir verwenden dabei die (binäre) Entropie $H$.

#### Binäre Entropy

Die binäre Entropy von einem Node kann man folgendermassen berechnen:

\begin{align}
H(p) = -p \log_2(p) - (1 - p) \log_2(1 - p)
\end{align}

Dabei entsprich $p$ dem Anteil der positiven Klasse (Anteil Datenpunkte mit $y=1$)

Man kann leicht verifizieren, dass die binäre Entropie minimal wird wenn alle Datenpunkte von derselben Klasse sind (Achtung: $p\log_2(p)=0$ falls $p=0$.)



#### Information Gain

Ganz spezifisch möchte man _Information Gain_ $IG$ optimieren. Das ist die Differenz in der _Impurity_ $Q$ bevor und nach einem Split.

Formal wird der Information Gain für einen Split $S_{j,t}$ auf einem Datensatz $\mathbf{X}$ berechnet als:

\begin{equation*}
IG(\mathbf{X}, S_{j,t}) = Q(\mathbf{X}) - \Big( \frac{|\mathbf{X}_{l}|}{|\mathbf{X}|} Q(\mathbf{X}_l) + \frac{|D_{r}|}{|\mathbf{X}|}  Q(\mathbf{X}_r) \Big)
\end{equation*}

Dabei löst man an jedem Node das Optimisierungsproblem:

\begin{equation*}
\operatorname*{argmax}_{j,t} IG(\mathbf{X}, S_{j,t})
\end{equation*}

Die _Impurity_ wird also noch gewichtet mit den Anzahl Beobachtungen in den Child-Nodes. Bei binären Splits gibt es immer genau zwei Child-nodes, einen linken $\mathbf{X}_l$ und einen rechten $\mathbf{X}_r$.


#### Stopping Criteria

Man stoppt den rekursiven Algorithmus und erstellt ein _Leaf Node_ wenn eine der folgende Bedingungen eintreffen:

- Man hat eine vordefinierte maximale Tiefe erreicht (wobei der Root-Node auf Tiefe 0 ist).
- Alle Beobachtungen in $\mathbf{X}$ sind von derselben Klasse.
- $\operatorname*{argmax}_{j,t} IG(\mathbf{X}, S_{j,t})$ ist kleiner als ein festgelegter Threshold $\epsilon$
- Die Anzahl Beobachtungen in einem Node ist kleiner als ein festgelegter Threshold $m$
- Jeder Leaf-Node muss mindestens eine Beobachtung haben.


### Notation

- $n$: Anzahl Beobachtungen in einem Datensatz
- $k$: Anzahl Features in einem Datensatz
- $x^{(i)}_j$: Wert von Feature $j$ von Beobachtung $\mathbf{x}^{(i)}$
- $S_{j,t}$: Split-Kriterium. Dieses splittet entlang von Feature $j$ mit Threshold $t$. Beobachtungen mit $x_j<t$ gehen nach links, Beobachtungen mit $x_j \ge t$ nach rechts.
- $t$: Threshold für Feature-Split $\in (-\infty , +\infty)$
- $j$: Feature Index $\in [0, k-1]$
- $p$: Anteil positive Beobachtungen / Wahrscheinlichkeit dass eine Beobachtung positiv ($y=1$) ist $\in [0, 1]$.
- $y$: Label einer Beobachtung $\in \{0, 1\}$
- $\mathbf{X}$: Ein Datensatz. Alle Datenpunkte / Beobachtungen **von einem bestimmten** Node (ausser beim Root-Node sind es also nicht alle Datenpunkte).
- $IG(\mathbf{X}, S_{j,t})$: Information Gain von $\mathbf{X}$, falls man mit $S_{j,t}$ splittet.
- $Q(\mathbf{X})$: Die Impurity von einem Datensatz $\mathbf{X}$.
- $H(p)$ Die Entropie von $p$.
- $|\mathbf{X}|$ die Anzahl Beobachtungen in $\mathbf{X}$.
- $|\mathbf{X}_l|$ die Anzahl Beobachtungen im linken Child-Node nach einem Split von $\mathbf{X}$ (Beobachtungen mit $x_j<t$)
- $|\mathbf{X}_r|$ die Anzahl Beobachtungen im rechten Child-Node nach einem Split von $\mathbf{X}$ (Beobachtungen mit $x_j \ge t$)
- $\log_2(x)$: Logarithmus von $x$ zur Basis $2$.

### Aufbabe 2a

Vervollständigen Sie die folgenden Klassen und Funktionen.

<!-- grade_id: answer2a, points: 5, grade: true, solution: true -->
```python
from typing import Callable

import numpy as np


def binary_entropy(p: np.ndarray) -> np.ndarray:
    """Calculates the binary entropy over an array of positive proportions
    Args:
        p: (z,) 1-dimensional array with z elements, 
            each element denotes the proportion of positive samples
    Returns:
        (z,) the binary entropy of each proportion
    """
    # YOUR CODE HERE
    raise NotImplementedError()


class BinarySplitOptimizer:
    """Finds the best possible binary split of a dataset X, given a specific binary impurity function Q."""
    
    def __init__(self, impurity: Callable[np.ndarray, np.ndarray]):
        """
        Args:
            impurity: A function/callable that calculates the impurity given probabilities.
                takes as input a 1-D array of k probabilities (z, )
                returns z impurity values  (z, )
        """
        self.impurity = impurity
    
    def find_optimal_split(self, X: np.ndarray, y: np.ndarray) -> (int, float, float):
        """
        Finds the best split in X to minimize the impurity Q of the child-nodes.

        Args: 
            X (n, k): Input features, where n is the number of samples and k is the number of features.
            y (n, ): Target values.
        
        Returns:
            Tuple(j, t, impurity)
            
                j: The id of the feature 'j' to split, in  [0, k-1]
                t: the threshold 't' of the feature to split
                impurity: the impurity value
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def find_optimal_threshold(self, x: np.ndarray, y: np.ndarray) -> (float, float):
        """Find the threshold 't' that minimizes the impurity Q for a given feature vector x.

        Args:
            x: (n, ) The value of a specific feature for each of n data points
            y: (n, ) The target / class label for each of n data points

        Returns:
            Tuple(t, impurity)
            Threshold 't' that minimizes the impurity.
        """
        # YOUR CODE HERE
        raise NotImplementedError()

```

### Aufgabe 2b

Die folgende Zelle enthält verschiedene Tests die Ihre Implementation prüfen. Sorgen Sie dafür, dass die folgenden Tests erfolgreich sind. Stellen Sie sicher, dass die Input-Shapes der Methoden die Sie implementieren den Doc-Strings entsprechen.

**Achtung: Die Tests decken nicht alles ab. Sie können also nicht davon ausegehen, dass Ihre Implementation korrekt ist sobald die Tests erfolgreich sind.**

Es ist grundsätzlich ihre Aufgabe, die Implementation genau zu prüfen. Sie können dazu weitere Zellen mit eigenen Tests einfügen. Sie können jedoch die folgende Zelle nicht ändern. Diese wird nach Abgabe wieder überschrieben, sodass die von mir definierten Tests ausgeführt werden.

<!-- grade_id: answer2b, points: 0, grade: true, solution: false -->
```python
import numpy as np

def print_result(test_name, passed, expected, actual):
    status = "Passed" if passed else "Failed"
    print(f"{status} test: {test_name}")
    print(f"----> Expected: {expected}")
    print(f"----> Actual: {actual}")


def run_test_binary_entropy():
    probabilities = np.array([0.5, 1.0, 0.0])
    expected = np.array([1.0, 0.0, 0.0])
    actual = binary_entropy(probabilities)

    try:
        np.testing.assert_almost_equal(actual, expected, decimal=2)
        print_result("run_test_binary_entropy_calculator", True, expected, actual)
    except AssertionError:
        print_result("run_test_binary_entropy_calculator", False, expected, actual)

    
def run_test_impurity_optimizer():
    x = np.array([1, 2, 3, 4, 5])  # Feature values
    y = np.array([0, 0, 1, 1, 1])  # Binary classification labels
    expected_threshold = 2.5  # All thresholds between 2 and 3 are identical, we choose their mean
    expected_impurity = 0.0  # This is the expected value if binary entropy is the measure

    impurity_optimizer = BinarySplitOptimizer(binary_entropy)
    actual_threshold, actual_impurity = impurity_optimizer.find_optimal_threshold(x, y)

    try:
        np.testing.assert_almost_equal(expected_threshold, actual_threshold)
        np.testing.assert_almost_equal(expected_impurity, actual_impurity)
        print_result(
            "run_test_impurity_optimizer", True, 
            (expected_threshold, expected_impurity),
            (actual_threshold, actual_impurity))
    except AssertionError as e:
        print_result(
            "run_test_impurity_optimizer", False,
            (expected_threshold, expected_impurity),
            (actual_threshold, actual_impurity))


def run_test_impurity_splitter():
    
    X = np.array([
        [1, 2],
        [9, 3],
        [3, 10],
        [4, 18],
        [5, 20]
    ])
    
    y = np.array([0, 0, 1, 1, 1])  # Binary classification labels
    expected_feature_id = 1  # Assuming the second feature is the best for splitting
    expected_threshold = 6.5  # Assuming the best split is between 3 and 10 for the second feature
    expected_impurity = 0.0  # This is the expected value if binary entropy is the measure

    binary_split_optimizer = BinarySplitOptimizer(binary_entropy)
    actual_feature_id, actual_threshold, actual_impurity = binary_split_optimizer.find_optimal_split(X, y)

    try:
        np.testing.assert_almost_equal(expected_threshold, actual_threshold)
        np.testing.assert_almost_equal(expected_impurity , actual_impurity)
        np.testing.assert_equal(expected_feature_id, actual_feature_id)
        print_result(
            "run_test_impurity_splitter", True,
            (expected_feature_id, expected_threshold, expected_impurity),
            (actual_feature_id, actual_threshold, actual_impurity))
    except AssertionError as e:
        print_result(
            "run_test_impurity_splitter", False, 
            (expected_feature_id, expected_threshold, expected_impurity),
            (actual_feature_id, actual_threshold, actual_impurity))


for test in [run_test_binary_entropy, run_test_impurity_optimizer, run_test_impurity_splitter]:
    try:
        test()
    except Exception as e:
        print(f"Error during testing - test: {test} error: {e}")

```

### Aufgabe 2c
Jetzt werden Sie den `BinaryDecisionTreeClassifier` implementieren.

<!-- grade_id: answer2c, points: 6, grade: true, solution: true -->
```python
from dataclasses import dataclass
from typing import Self, Callable

import numpy as np
from sklearn.base import BaseEstimator, ClassifierMixin
from tqdm.notebook import tqdm


# Helper classes that can be used to construct the decision tree
@dataclass
class LeafNode:
    value: float # value of leaf node (positive class proportion)
    depth: int # depth of the node
    is_leaf_node: bool = True


@dataclass
class TreeNode:
    feature_index: int # feature on which to split 
    threshold: float # threshold on which to split 
    depth: int # depth of the node
    left: Self | LeafNode # the left child node
    right: Self | LeafNode # the right child node
    is_leaf_node: bool = False   


class BinaryDecisionTreeClassifier(BaseEstimator, ClassifierMixin):
    """Binary Decision Tree Classifier
    
    Args:
    -----
        max_depth: max depth of the decision tree, 
            if max depth is reached a leaf node is created
        min_information_gain_to_split: minimum information gain required to split a node, 
            otherwise a leaf node is created
        min_samples_in_node_for_split: minimum number of samples required to split a node,
            otherwise a leaf node is created
        p_feature_sampling: proportion of features to consider at each node (random subsample), (0, 1]
        impurity: function / callable to calculate binary impurity
        random_seed: ensure reproducibility for random steps
        verbose: bool, whether to print log messages (optional)
    """
    
    def __init__(
        self,
        max_depth: int | None = None,
        min_information_gain_to_split = 0.0,
        min_samples_in_node_for_split: int = 2,
        p_feature_sampling: float = 1.0,
        impurity: Callable[np.ndarray, np.ndarray]=binary_entropy,
        random_seed: int=123,
        verbose: bool = False):

        self.max_depth = max_depth
        self.min_information_gain_to_split = min_information_gain_to_split
        self.min_samples_in_node_for_split = min_samples_in_node_for_split
        self.p_feature_sampling = p_feature_sampling
        self.impurity = impurity
        self.verbose = verbose
        self.random_seed = random_seed

        self._splitter = BinarySplitOptimizer(self.impurity)
        self._rng = np.random.default_rng(random_seed)
        
    def fit(self, X: np.ndarray, y: np.ndarray) -> Self:
        """Fit Decision Tree

        Args:
        -----
            X: (n, k)
            y: (n, )

        Returns:
        --------
            self
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def predict(self, X: np.ndarray) -> np.ndarray:
        """Return most likely class per sample

        Args:
        -----
            X: (n, k)

        Returns:
        --------
            (n, )
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def predict_proba(self, X: np.ndarray) -> np.ndarray:
        """Return probability of being the positive class per sample.

        Args:
        -----
            X: (n, k)

        Returns:
        --------
            (n, )
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def score(self, X: np.ndarray, y: np.ndarray) -> float:
        """Calculate Accuracy

        Args:
        -----
            X: (n, k)
            y: (n, )

        Returns:
        --------
            accuracy (float)
        """
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # any helper methods
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Aufgabe 2d

Nun überprüfen Sie grafisch ob Ihre Implementation funktioniert.

Fitten Sie dazu jeweils 2 Modelle auf jedem Beispiel-Datensatz.

Für jeden Modell-Fit und Datensatz: Zeigen Sie den Datensatz (Plot) und plotten Sie die _Decision Regions_ der Modelle. Verwenden Sie [mlxtend.plotting.plot_decision_regions](https://rasbt.github.io/mlxtend/api_subpackages/mlxtend.plotting/#plot_decision_regions).

Berechnen Sie jeweils die Accuracy mit `score(X, y)` und zeigen Sie diese oberhalb von den Decision-Region-Plots. Z.B mit: 

```
accuracy = dt.score(X, y)
_ = ax.set_title(f"Accuracy: {accuracy:.2f}")
```

Sie sollen also pro Datensatz jeweils 3 Plots zeigen (Daten, Decision-Region Modell 1, Decision-Region Modell 2).

Verwenden Sie folgende Parameter für Modell 1:

- `max_depth=3`
- `min_information_gain_to_split=0.0`
- `min_samples_in_node_for_split=3`
- `p_feature_sampling=1.0`


Verwenden Sie folgende Parameter für Modell 2:

- `max_depth=10`
- `min_information_gain_to_split=0.0`
- `min_samples_in_node_for_split=0`
- `p_feature_sampling=1.0`



<!-- grade_id: answer2d_a, points: 1, grade: true, solution: true -->
```python
from sklearn.datasets import make_classification, make_blobs, make_moons
import mlxtend
import seaborn as sns
from mlxtend.plotting import plot_decision_regions
import matplotlib.pyplot as plt
import numpy as np


def generate_ds1():
    # Number of observations in each class
    num_obs = 50
    rng = np.random.default_rng(123)

    # Generating x1 and x2
    x1 = np.linspace(-4, 4, num_obs).reshape(-1, 1)
    x2 = np.linspace(-4, 4, num_obs).reshape(-1, 1)

    # Initial binary classification, without noise
    X = np.hstack([x1, x2])
    X = np.vstack([X, X])  # Doubling for both classes
    class_perturbation = np.hstack([1 + rng.standard_normal(num_obs) * 0.2, -1 + rng.standard_normal(num_obs) * 0.2])
    X[:, 0] += class_perturbation  # Adding noise to the x1 component
    
    # Generating labels
    y = np.hstack([np.repeat(0, num_obs), np.repeat(1, num_obs)])
    return X, y


def generate_ds2():
    X, y = make_classification(
        n_samples=100, n_features=2, n_informative=1, n_redundant=0, n_repeated=0, 
        n_classes=2, n_clusters_per_class=1, weights=None, flip_y=0, class_sep=2.0, 
        hypercube=True, shift=0.0, scale=1.0, shuffle=True, random_state=123)
    return X, y


def generate_ds3():
    X, y = make_moons(noise=0.3, random_state=0)
    return X, y


def generate_ds4():
    # Defining centers for the clusters
    centers = [(-3, -3), (3, 3)]
    cluster_std = 0.5  # Standard deviation of the clusters

    # Generating the dataset
    X, y = make_blobs(n_samples=400, centers=centers, cluster_std=cluster_std, n_features=2, random_state=123)
    y = y % 2
    
    # add additional points
    X_add = np.array([
        [-3, 4],
        [0, -2],
        [4, -2],
        [-3, 2]
    ]).reshape(-1, 2)
    y_add = np.array([1, 1, 0, 0]).reshape(-1)
    
    X = np.vstack([X, X_add])
    y = np.hstack([y, y_add])

    return X, y


def generate_ds5():
    grid_size = 5
    rng = np.random.default_rng(seed=123)
    # Create grid points
    X = [(x, y) for x in range(grid_size) for y in range(grid_size)]
    # Assign classes
    y = [(x + y) % 2 for x, y in X]
    #X = [(x - (grid_size % 2), y - (grid_size % 2)) for x, y in X]
    X = [(x + rng.uniform(-0.1, 0.1), y + rng.uniform(-0.1, 0.1)) for x, y in X]
    return np.array(X), np.array(y)


# iterating over the toy datasets
for i, (X, y) in enumerate([generate_ds1(), generate_ds2(), generate_ds3(), generate_ds4(), generate_ds5()]):
    
    # Plot the dataset
    fig, ax = plt.subplots(figsize=(12, 4), ncols=3)
    _ = sns.scatterplot(x=X[:, 0], y=X[:, 1], hue=y, ax=ax[0]).set(title=f"Dataset {i + 1}")

    # YOUR CODE HERE
    raise NotImplementedError()
    
    plt.tight_layout()

    

```

Erläutern Sie die Ergebnisse von jedem Datensatz. Argumentieren Sie klar und nachvollziehbar indem Sie auf die Grafiken verweisen.
- Diskutieren Sie kurz die Güte der Modellierung. Welches Modell (1 oder 2) gefällt ihnen besser und warum?
- Falls Unterschiede zwischen den Modellen sichtbar sind: Erklären Sie diese aufgrund der Hyperparameter und reflektieren Sie ob ihre Implementation wie erwartet funktioniert.

<!-- grade_id: answer2d_b, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

### Aufgabe 2e
Nun implementieren Sie den `BinaryRandomForestClassifier`.

<!-- grade_id: answer2e, points: 3, grade: true, solution: true -->
```python
from typing import Self
from tqdm.notebook import tqdm

from sklearn.base import ClassifierMixin, BaseEstimator


class BinaryRandomForestClassifier(BaseEstimator, ClassifierMixin):
    """Binary Random Forest Classifier
    
    Args:
    -----
        num_trees: the number of decision trees to fit
        max_depth: max depth of the decision tree, 
            if max depth is reached a leaf node is created
        min_information_gain_to_split: minimum information gain required to split a node, 
            otherwise a leaf node is created
        min_samples_in_node_for_split: minimum number of samples required to split a node,
            otherwise a leaf node is created
        p_feature_sampling: proportion of features to consider at each node (random subsample), (0, 1]
        impurity: function / callable to calculate binary impurity
        random_seed: ensure reproducibility for random operations
        verbose: bool, whether to print log messages (optional)
    """
    def __init__(
        self, 
        num_trees: int = 10,
        max_depth: int | None = None,
        min_information_gain_to_split: float = 0.0,
        min_samples_in_node_for_split: int = 2,
        p_feature_sampling: float = 1.0,
        verbose: bool=False,
        random_seed: int=123
    ):
        self.num_trees = num_trees
        self.max_depth = max_depth
        self.min_samples_in_node_for_split = min_samples_in_node_for_split
        self.min_information_gain_to_split = min_information_gain_to_split
        self.p_feature_sampling = p_feature_sampling
        self.random_seed = random_seed
        self.verbose = verbose
        self._trees = []

    def fit(self, X: np.ndarray, y: np.ndarray) -> Self:
        """Fit the model

        Args:
        -----
            X: (n, k)
            y: (n, )

        Returns:
        --------
            self
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def predict(self, X: np.ndarray) -> np.ndarray:
        """Return most likely class per sample

        Args:
        -----
            X: (n, k)

        Returns:
        --------
            (n, )
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def predict_proba(self, X: np.ndarray) -> np.ndarray:
        """Return probability of being the positive class per sample.

        Args:
        -----
            X: (n, k)

        Returns:
        --------
            (n, )
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def score(self, X: np.ndarray, y: np.ndarray) -> float:
        """Calculate Accuracy

        Args:
        -----
            X: (n, k)
            y: (n, )

        Returns:
        --------
            a
        """
        # YOUR CODE HERE
        raise NotImplementedError()
    # any helper methods if needed
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Aufgabe 2f

Nun testen Sie den `BinaryRandomForestClassifier`. Verwenden Sie denselben Ansatz wie aus Aufgabe 2d. Trainieren Sie pro Datensatz jeweils zwei Modelle mit folgenden Parametern:

Verwenden Sie folgende Parameter für Modell 1:

- `num_trees=5`
- `max_depth=10`
- `min_information_gain_to_split=0.0`
- `min_samples_in_node_for_split=0`
- `p_feature_sampling=1.0`



Verwenden Sie folgende Parameter für Modell 2:

- `num_trees=30`
- `max_depth=10`
- `min_information_gain_to_split=0.0`
- `min_samples_in_node_for_split=0`
- `p_feature_sampling=1.0`

<!-- grade_id: answer2f_a, points: 1, grade: true, solution: true -->
```python
# iterating over the toy datasets
for i, (X, y) in enumerate([generate_ds1(), generate_ds2(), generate_ds3()]):
    
    # Plot the dataset
    fig, ax = plt.subplots(figsize=(12, 4), ncols=3)
    _ = sns.scatterplot(x=X[:, 0], y=X[:, 1], hue=y, ax=ax[0]).set(title=f"Dataset {i + 1}")

    # YOUR CODE HERE
    raise NotImplementedError()
    
    plt.tight_layout()


```

Erläutern Sie die Ergebnisse von jedem Datensatz. Argumentieren Sie klar und nachvollziehbar indem Sie auf die Grafiken verweisen.
- Diskutieren Sie die Güte der Modellierung. Welches Modell (1 oder 2) gefällt ihnen besser und warum?
- Falls Unterschiede zwischen den Modellen sichtbar sind: Erklären Sie diese aufgrund der Hyperparameter und reflektieren Sie ob ihre Implementation wie erwartet funktioniert.
- Erwähnen und diskutieren Sie Unterschiede zu den Tests aus 2d. Inwiefern sieht man, dass es sich hier um einen RandomForest handelt?

<!-- grade_id: answer2f_b, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Aufgabe 3 (12 Punkte)

In dieser Aufgabe modellieren Sie die Hotel-Daten und versuchen ein möglichst gutes Modell zu finden.

Hinweis: Sollte es Ihnen nicht gelungen sein die Modelle zu implementieren können Sie [sklearn.ensemble.RandomForestClassifier](https://scikit-learn.org/stable/modules/generated/sklearn.ensemble.RandomForestClassifier.html) und [sklearn.tree.DecisionTreeClassifier](https://scikit-learn.org/stable/modules/generated/sklearn.tree.DecisionTreeClassifier.html) verwenden.

### Aufgabe 3a

Trainieren Sie verschiedene Modelle der Klassen `BinaryRandomForestClassifier` und `BinaryDecisionTreeClassifier` auf dem Datensatz `hotel_train.csv` um `IstStorniert` vorherzusagen.

Vergleichen Sie die Modelle und finden Sie die beste Konfiguration. Verwenden Sie Kreuzvalidierung.

Berechnen Sie den Score vom besten Modell auf dem Trainings-Datensatz und auf dem Validierungs-Datensatz `hotel_validation.csv` und geben Sie diese aus.

<!-- grade_id: antwort3a, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

### Aufgabe 3b

Interpretieren Sie das Ergebnis:

- Welche Modell-Varianten haben Sie verglichen? Begründen Sie kurz warum Sie welche Parameter wie variiert haben.
- Vergleichen und interpretieren Sie Scores auf Trainings- und Validation-Datensatz.

<!-- grade_id: answer3b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Aufgabe 3c

Beantworten Sie folgende Fragen:

- Mika fragt Sie ob es hilfreich wäre mehr Trainingsdaten zu sammeln (was viel Zeit und Geld kosten würde). Kann man diese Frage beantworten? Falls ja, wie würden Sie vorgehen?
- Mika muss eine Präsentation für das Hotel-Management vorbereiten und die Ergebnisse, sowie Empfehlungen vorstellen. Mika bittet Sie um Rat. Wie könnte das Management dieses Modell nun gewinnbringend einsetzen?

<!-- grade_id: answer3c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Aufgabe 4 (Bonusaufgabe, 5 Punkte) Competition

In dieser Aufgabe geht es darum auf einem Testset Vorhersagen zu generieren die möglichst akkurat sind. Sie dürfen dazu das Modell aus Aufgabe 3 verwenden oder ein komplett neues Modell erstellen. Sie dürfen dabei alle Methoden aus `sklearn` verwenden (auch andere Modelle und Algorithmen).

Sie müssen die Vorhersagen abspeichern und mit dem finalen Commit abgeben. 

Ihre Vorhersagen werden dann beurteilt und mit denjenigen Ihrer Kolleginnen und Kollegen verglichen.

Die besten kriegen die Bonuspunke.

Beachten Sie die maximale Laufzeit vom Notebook von 30 Minuten.

### Aufgabe 4a

Lesen Sie das File `hotel_test.csv` ein und generieren Sie die Vorhersagen.  Speichern Sie diese in das File `hotel_test_predictions.csv`. 

Definieren Sie im Code die Variable `my_leaderboard_name`. Dies ist Ihr Pseudonym für die öffentliche Anzeige von ihrem Score.

<!-- grade_id: answer4a, points: 5, grade: true, solution: true -->
```python
from pathlib import Path

from sklearn.metrics import accuracy_score

# Read the test dataset
df_test = pd.read_csv("hotel_test.csv")

# TODO: Change `my_leaderboard_name` to a name of your liking
my_leaderboard_name = 'dummy'

# TODO: these are dummy predictions to demonstrate how the output should look like.
# replace them with better predictions!

fname_predictions = "hotel_test_predictions.csv"
y_hat = np.array([0 for _ in range(df_test.shape[0])]).reshape(-1)
df_test_predictions = pd.DataFrame({"IstStorniert": y_hat})                                             
df_test_predictions.to_csv(fname_predictions, index=False)

# YOUR CODE HERE
raise NotImplementedError()
```

Der folgende Code soll / kann nicht verändert werden. Dieser liest ihre Vorhersagen ein, evaluiert diese, und schreibt das Resultat in ein File.

```python
# ensure that `hotel_test_predictions.csv` exists!
if Path(fname_predictions).exists():
    has_test_predictions = True
    print(f"Reading Predictions from {fname_predictions}")
    df_test_predictions = pd.read_csv(fname_predictions, index_col=False)
    y_hat = df_test_predictions[["IstStorniert"]].to_numpy().reshape(-1)
else:
    has_test_predictions = False
    print(f"File {fname_predictions} not found")

# reads the y_true for the test dataset (not available to you!)
path_df_test_solutions = "/home/jovyan/work/data/hotel_bookings/hotel_test_solution.csv"

if Path(path_df_test_solutions).exists() & has_test_predictions:
    print(f"File {path_df_test_solutions} found")
    df_test_true = pd.read_csv(path_df_test_solutions, index_col=False)
    y_true = df_test_true[["IstStorniert"]].to_numpy().reshape(-1)
    print(f"Test Accuracy: {accuracy_score(y_true, y_hat)}")

    with open("test_accuracy.txt", "w") as f:
        _ = f.write(f"{accuracy_score(y_true, y_hat):.5f},{my_leaderboard_name}")
else:
    print(f"File {path_df_test_solutions} not found")
```