## FHNW Machine Learning - FS2024

```python
# DO NOT EDIT THIS CELL
# ADD IMOPRTS IN ASSIGNMENT CELLS

# Notebook configs
%load_ext autoreload
%autoreload 2
%matplotlib inline

from IPython.core.interactiveshell import InteractiveShell
from IPython.display import HTML

InteractiveShell.ast_node_interactivity = "all"
```

# Bewertete Aufgabe 01

**Ausgabe:** Montag 4.3.2024

**Abgabe:** Montag 25.3.2024 23:59 Uhr


## Vorgaben zu Umsetzung und Abgabe

- Die Algorithmen müssen auf der Basis von Array Operationen selber implementiert werden.
- Der Code muss lauffähig sein bei Ausführung im Docker-Container zum Repo oder auf JHub. 
- Es darf kein Code ausgelagert werden, i.e. sämtlicher Code muss sich im Notebook befinden.
- Sämtliche Plots sind komplett beschriftet (Achsen, Labels, Überschrift, Colorbar, ..), so dass der Plot ohne den Code zu konsultieren, verstanden werden kann.
- Als **Abgabe** zählt der letzte Commit vor Abgabetermin in in Ihrem Fork des Repos.  

- **Bitte löschen, kopieren, duplizieren, splitten und verschieben Sie die vorhandenen Zellen nicht**. Dies führt zu Problemen bei der Korrektur. Sie dürfen aber beliebig viele weitere Zellen hinzufügen (nur via **insert new cell**).
- Laufzeit vom Notebook: Das Notebook sollte in weniger als 30 Minuten ausgeführt werden können.

Für die Erarbeitung der Lösung darf unter Studierenden zusammengearbeitet werden. Die Zusammenarbeit ist dabei aber auf konzeptionelle und algorithmische Fragen und Verständnisaspekte beschränkt.  

**Es darf kein Code oder Text von anderen oder vom Internet kopiert werden.**


### Module

Neben den Python-Basismodulen dürfen Sie die folgenden Module immer benutzen: `numpy`, `pandas`, `matplotlib`, `seaborn`,  `tqdm` (für Progress-Bars).

Sie dürfen auch generell [sklearn.preprocessing](https://scikit-learn.org/stable/modules/classes.html#module-sklearn.preprocessing), [sklearn.model_selection](https://scikit-learn.org/stable/model_selection.html), [sklearn.pipeline](https://scikit-learn.org/stable/modules/classes.html#module-sklearn.pipeline) und [sklearn.compose](https://scikit-learn.org/stable/modules/classes.html#module-sklearn.compose) benutzen.

Zusätzliche Module dürfen Sie nur verwenden wenn ausdrücklich erwähnt oder bereits vorgegeben in der Code-Cell.

## Bewertung

Bewertet wird:

- Vollständigkeit (Code, Text)
- Korrektheit (Code, Text)
- Implementation (z.B. Vektorisierung der Operationen, Scikit-Learn API, Visualisierungen, Lesbarkeit Code/Output)

## Einleitung

Der Schweizer Bauernverband hat immer mehr Bäuerinnen und Bauern welche Heidelbeeren anbauen. Mittlerweilen sind es schweizweit ca 150 Fussballfelder. Der Bedarf an Heidelbeeren ist stark gewachsen, man kommt mit der Produktion kaum nach. Es gibt dabei starke Konkurrenz aus dem Ausland. Der Bauernverband und viele Konsumentinnen und Konsumenten möchten aber schweizer Heidelbeeren auf dem Tisch. Der Bauernverband hat für die Heidelbeerproduzenten eine Untersuchung gestartet um herauszufinden, welche Faktoren die Ernte beeinflussen. Man möchte die Produzenten dabei unterstützen die Ernte zu maximieren und auch die voraussichtliche Ernte vorherzusagen. 

Der Bauernverband hat bei verschiedenen Produzenten über die vergangenen Jahre Daten erhoben hat. Nun hat der Verband Alex beauftragt ein Modell zu erstellen um den Zusammenhang zwischen Ernte und verschiedenen Faktoren zu modellieren. Mit diesem Modell erhofft man sich weitere Erkenntnisse. In einem ersten Schritt geht es darum eine möglichst akkurate Vorhersage der Ernte zu ermöglichen.

Die folgende Tabelle zeigt die vorhandenen Daten. Pro Heidelbeer-Feld hat man folgende Daten über die Saison erhoben:


| Spalte        | Erklärung                                                      |
|-----------------------------|----------------------------------------------------------------|
| Id              | Eindeutige Identifikationsnummer für jede Beobachtung              |
| Klongrösse                   | Die durchschnittliche Grösse der Heidelbeerkolonie im Feld  ($m^2$)      |
| Honigbiene                  | Dichte der Honigbienen im Feld (Bienen pro Quadratmeter pro Minute) |
| Hummeln                     | Dichte der Hummeln im Feld (Hummeln pro Quadratmeter pro Minute)    |
| AvgTempHoch | Durchschnitt der täglichen maximalen Lufttemperatur während der Blütezeit (C°)    |
| AvgTief | Durchschnitt der täglichen minimalen Lufttemperatur während der Blütezeit (C°)   |
| Regentage                   | Gesamtzahl der Tage während der Blütezeit mit Niederschlag      |
| Fruchtansatz                | Prozentsatz der Blüten, die zu Früchten werden (0-1)                 |
| Fruchtmasse                 | Gewicht der einzelnen Frucht (g)                                  |
| Samen                       | Anzahl der Samen pro Frucht                                    |
| Ertrag                      | Gesamtertrag der Heidelbeerernte (kg/ha)                              |



Die Aufgabe besteht darin Alex zu unterstützen ein Modell zu entwickeln mit dem der `Ertrag` vorhergesagt werden kann.

## Aufgabe 1 (6 Punkte)

Bevor man einen Datensatz modelliert ist es wichtig diesen zu untersuchen und kennenzulernen. Benutzen Sie ihre Kenntnisse aus _dsp_ um die Daten einzulesen, zu untersuchen und zu visualisieren.

### Aufgabe 1 a

Unternehmen Sie die folgenden Schritte:
    
- Lesen Sie den Datensatz `heidelbeere_train.csv` ein und verschaffen Sie sich einen groben Überblick.
- Berechnen und zeigen Sie wichtige Kennzahlen.
- Visualisieren Sie jede Variable mit einem sinnvollen Plot. Berücksichtigen Sie dabei die Fragestellung.


Versetzen Sie sich in Ihre Lage: Sie kennen den Datensatz nicht und müssen nun so viele Erkenntnisse wie nötig erhalten, damit Sie die Daten verstehen und modellieren können. Ausserdem möchten Sie abschätzen können welche Variablen wichtig sein könnten und wie man Variablen allenfalls prozessieren muss.

<!-- grade_id: antwort1a, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Aufgabe 1 b

Unser Ziel ist es, die Variable `Ertrag` vorherzusagen unter Verwendung der anderen Variablen.  

Was lässt sich über dieses Vorhaben sagen? Diskutieren Sie Ihre Einsichten auf Basis der Analyse in Aufgabe 1a.

Was fällt sonst noch auf in den Daten? Beschreiben und begründen Sie.

<!-- grade_id: antwort1b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Aufgabe 2 (8 Punkte)

In dieser Aufgabe ist es das Ziel eine regularisierte Lineare Regression zu implementieren (_Ridge Regression_) und die Implementation zu validieren.

### Ridge Regression

Ridge Regression ist eine regularisierte Form ($L_2$-Regularisierung) der Ordinary Least Squares (OLS) Kostenfunktion für die lineare Regression.  

Die Ridge Regression-Kostenfunktion $J(\mathbf{w})$ für einen Datensatz $\cal{D}=\{(\mathbf{x}^{(1)}, y^{(1)}), \dots, (\mathbf{x}^{(n)}, y^{(n)}) \}$ mit $n$ Datenpunkten ist:

\begin{align}
J(\mathbf{w}) &= \frac{1}{2n}\sum_{i=1}^n (y^{(i)} - \mathbf{x}^{(i)T}\mathbf{w})^2 + \Omega(\mathbf{w})
\end{align}

Wobei die Regularisierung $\Omega(\mathbf{w})$ folgendermassen definiert ist:

\begin{align}
\Omega(\mathbf{w}) &= \frac{1}{2}\lambda \sum_{j=1}^p w_j^2 
\end{align}

Der $i$te Datenpunkt $\mathbf{x}^{(i)}$ ist ein Vektor der Dimensionalität $\mathbb{R}^{p +1}$: $\mathbf{x}^{(i)} = \Big(1, x_1^{(i)}, \dots , x_p^{(i)}\Big)$.

$\mathbf{w} = (w_0, \dots, w_p)$ sind dabei die Modellkoeffizienten, $\lambda$ ist die Regularisierungsstärke.


**Beachten Sie:**

- In sklearn wird statt $\lambda$ jeweils $\alpha$ `alpha` als Bezeichnung für die Regularisierungsstärke verwendet (wohl weil `lambda` ein reserviertes Wort ist in Python).

- Um Gradient Descent auszuführen müssen Sie die Kostenfunktion ableiten. Die Vorgehensweise ist in den Kursunterlagen erklärt.

- Coursera verwendet für die Bezeichung der Anzahl Datenpunkte $m$. Hier verwenden wir $n$, was gebräuchlicher ist. Ausserdem: $w$ ist ein Skalar, $\mathbf{w}$ ein Vektor und $\mathbf{W}$ eine Matrix.

- Implementieren Sie alles vektorisiert.


### Stochastic Gradient Descent

Mit _Batch Gradient Descent (BGD)_ berechnet man den Gradienten der Kostenfunktion bezüglich der Modellparameter für jeden _Gradient Descent Step_ mit **allen Datenpunkten**. Damit berechnet man den Gradient exakt. Man kann den Gradient jedoch auch mit einem Subset der Datenpunkte berechnen und den exakten Gradient damit schätzen. Dadurch verringert sich der Rechenaufwand. Mit _Stochastic Gradient Descent (SGD)_ nimmt man nur **einen, zufällig ausgewählten, Datenpunkt** (deshalb _stochastic_) um den Gradient in jeder Iteration zu schätzen. 

Implementieren Sie die Optionen `optimization_method="bgd"` und `optimization_method="sgd"` um das Modell wahlweise mit _BGD_ oder _SGD_ zu optimisieren. 

### Aufgabe 2a
Ergänzen Sie die untenstehende Klasse und deren Methoden. Folgen Sie der Scikit-Learn API ([Link](https://scikit-learn.org/stable/developers/develop.html)): Das bedeutet im Wesentlichen, dass Sie die Methoden implementieren sollen, die in der Klasse schon vorgegeben sind.

Beachten Sie die in den Doc-Strings spezifizierten Angaben, insbesondere auch die der Shapes der Inputs und Outputs der einzelnen Methoden.

Der Estimator soll eine Ridge Regression durchführen. Das Finden der Modell-Koeffizienten soll mit Gradient Descent erfolgen.

<!-- grade_id: antwort2a, points: 8, grade: true, solution: true -->
```python
from typing import Self

import numpy as np
from sklearn.base import BaseEstimator
from tqdm.notebook import tqdm


class LinearRegression(BaseEstimator):
    """Linear Regression
    
    Args:
    -----
        epsilon: if norm of gradient falls below epsilon, 
            gradient descent terminates (disable with negative values)  
        max_num_steps: max number of steps for gradient descent
        learning_rate: learning rate for gradient descent
        optimization_method: one of 'bgd' (batch gradient descent),
            'sgd' (stochastic gradient descent)
        alpha: regularization strength (lambda)
        verbose: whether to print progress during model training (optional)
    """

    def __init__(
        self,
        epsilon: float = -1,
        max_num_steps: int = 1000,
        learning_rate: float = 0.1,
        optimization_method: str = "bgd",
        alpha: float = 0.0,
        verbose: bool = True
    ):
        self.alpha = alpha
        self.epsilon = epsilon
        self.max_num_steps = max_num_steps
        self.learning_rate = learning_rate
        self.optimization_method = optimization_method
        self.verbose = verbose

    def fit(self, X: np.ndarray, y: np.ndarray) -> Self:
        """Fit the model coefficients.
        Args:
            X: input data (n, p)
            y: input labels (n, )

        Returns:
            self
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def gradient(self, X: np.ndarray, y: np.ndarray) -> np.ndarray:
        """Calculate Gradient of Cost Function.
        
        Calculates dJ/dw, while w are the current parameter estimates in self.w_

        Args:
            X: input data (n, p)
            y: input labels (n, )

        Returns:
            dJ/dw gradient vector (p + 1, )
        """
        # YOUR CODE HERE
        raise NotImplementedError()
        assert gradient.ndim == 1                
        return gradient
        
    def cost(self, X: np.ndarray, y: np.ndarray) -> float:
        """Evaluate the Cost Function.

        Args:
            X: input data (n, p)
            y: input labels (n, )

        Returns:
            total cost
        """
        # YOUR CODE HERE
        raise NotImplementedError()
        assert isinstance(cost, float)
        return cost

    def predict(self, X: np.ndarray) -> np.ndarray:
        """Calculate Predictions.

        Args:
            X: input Data (n, p)
            
        Returns:
            Predictions (n, )
        """
        # YOUR CODE HERE
        raise NotImplementedError()
        assert predictions.ndim == 1
        return predictions

    def score(self, X: np.ndarray, y: np.ndarray) -> float:
        """Calculate R^2
        
        See: https://en.wikipedia.org/wiki/Coefficient_of_determination

        Args:
            X: Input Data (n, p)
            y: Input Labels (n, )
        
        Returns:
            R^2
        """
        # YOUR CODE HERE
        raise NotImplementedError()
        assert isinstance(score, float)
        return score

    # Other methods if needed
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Aufgabe 2b
Die folgende Zelle enthält verschiedene Tests die Ihre Implementation prüfen. Sorgen Sie dafür, dass die folgenden Tests erfolgreich sind. Stellen Sie sicher, dass die Input-Shapes der Methoden die Sie implementieren den Doc-Strings entsprechen.

**Achtung: Die Tests decken nicht alles ab. Sie können also nicht davon ausegehen, dass Ihre Implementation korrekt ist sobald die Tests erfolgreich sind.**

Ihre Abgabe wird noch mit weiteren, für Sie nicht sichtbare Tests, geprüft. Es ist grundsätzlich ihre Aufgabe, die Implementation genau zu prüfen. Sie können dazu weitere Zellen mit eigenen Tests einfügen. Sie können jedoch die folgende Zelle nicht ändern. Diese wird nach Abgabe wieder überschrieben, sodass die von mir definierten Tests ausgeführt werden.

<!-- grade_id: tests2b, points: 0, grade: true, solution: false -->
```python
import numpy as np
from sklearn.datasets import make_regression


def print_result(test_name, passed, expected, actual):
    status = "Passed" if passed else "Failed"
    print(f"{status} test: {test_name}")
    print(f"----> Expected: {expected}")
    print(f"----> Actual: {actual}")


def run_test_multivariate():
    Xdummy, ydummy, w_true = make_regression(
        n_samples=30, bias=5.0, coef=True, n_features=2, n_targets=1, random_state=123
    )
    w_true = np.concatenate([np.array([5.0]), w_true])

    lr = LinearRegression(alpha=0, epsilon=1e-3, optimization_method="bgd", learning_rate=0.1)
    lr.fit(Xdummy, ydummy)

    try:
        np.testing.assert_allclose(w_true, lr.w_, atol=1e-3)
        print_result("run_test_multivariate", True, w_true, lr.w_)
    except AssertionError:
        print_result("run_test_multivariate", False, w_true, lr.w_)


def run_test_univariate():
    Xdummy, ydummy, w_true = make_regression(
        n_samples=30, bias=5.0, coef=True, n_features=1, n_targets=1, random_state=123
    )
    w_true = np.concatenate([np.array([5.0]), np.array([w_true])])
    
    lr = LinearRegression(alpha=0, epsilon=1e-3, optimization_method="bgd", learning_rate=0.1)
    lr.fit(Xdummy, ydummy)
    try:
        np.testing.assert_allclose(w_true, lr.w_, atol=1e-3)
        print_result("run_test_univariate", True, w_true, lr.w_)
    except AssertionError:
        print_result("run_test_univariate", False, w_true, lr.w_)


def run_test_score():
    Xdummy, ydummy, w_true = make_regression(
        n_samples=30, bias=5.0, coef=True, n_features=1, n_targets=1, random_state=123
    )
    lr = LinearRegression(alpha=0, epsilon=1e-3, optimization_method="bgd", learning_rate=0.1)
    lr.fit(Xdummy, ydummy)
    score = lr.score(Xdummy, ydummy)
    expected_score = 1.0

    try:
        np.testing.assert_almost_equal(score, expected_score, decimal=3)
        print_result("run_test_score", True, expected_score, score)
    except AssertionError:
        print_result("run_test_score", False, expected_score, score)


def run_test_cost():
    Xdummy, ydummy, _ = make_regression(
        n_samples=30, bias=5.0, coef=True, n_features=1, n_targets=1, random_state=123
    )
    lr = LinearRegression(alpha=0, epsilon=1e-3, optimization_method="bgd", learning_rate=0.1)
    lr.fit(Xdummy, ydummy)
    actual_cost = lr.cost(Xdummy, ydummy)
    expected_cost = 1e-6

    try:
        assert actual_cost < expected_cost
        print_result("run_test_cost", True, f"< {expected_cost}", actual_cost)
    except AssertionError:
        print_result("run_test_cost", False, f"< {expected_cost}", actual_cost)


def run_test_gradient():

    Xdummy, ydummy, w_true = make_regression(
        n_samples=30, bias=5.0, coef=True, n_features=1, n_targets=1, random_state=123
    )
    w_true = np.concatenate([np.array([5.0]), np.array([w_true])])

    lr = LinearRegression(alpha=0, epsilon=1e-3, optimization_method="bgd", learning_rate=0.1)
    lr = lr.fit(Xdummy, ydummy)
    
    expected_gradient = np.array([-6.92655851, -59.01320294])
    weights = np.array([0.0, 0.0])
    lr.w_ = weights
    actual_gradient = lr.gradient(Xdummy, ydummy)
    try:
        np.testing.assert_allclose(actual_gradient, expected_gradient)
        print_result("run_test_gradient", True, expected_gradient, actual_gradient)
    except AssertionError:
        print_result("run_test_gradient", False, expected_gradient, actual_gradient)


for test in [run_test_multivariate, run_test_univariate, run_test_score, run_test_cost, run_test_gradient]:
    try:
        test()
    except Exception as e:
        print(f"Error during testing - test: {test} error: {e}")


```

## Aufgabe 3 (16 Punkte)

In dieser Aufgabe wenden Sie den Algorithmus auf dem Datensatz `heidelbeere_train.csv` an. Damit testen und untersuchen Sie ihre Implementation der Klasse `LinearRegression` und lernen gleichzeitig den Datensatz besser kennen.


### Aufgabe 3a

1) Trainieren Sie ein Modell mit `learning_rate=1.0` (der Rest soll auf den Default-Werten sein). Verwenden Sie ausschliesslich `Fruchtmasse` als Input-Variable um `Ertrag` vorherzusagen. Lassen Sie die Variablen komplett unverändert.
2) Berechnen und zeigen Sie $R^2$ mit der `score()` Methode.
3) Zeigen Sie in einer Tabelle für die ersten 10 Beobachtungen: `y_true` ($\mathbf{y}$), `y_hat` ($\mathbf{\hat{y}}$)und `y_true - y_hat` (das Residuum $\mathbf{r}$). Sie können ein `pd.DataFrame` erstellen und dieses mit `print` anzeigen.
4) Berechnen und zeigen Sie den Wert der Kostenfunktion.

Verwenden Sie `print()` Statements um Fragen nach bestimmten Outputs zu beantworten. Beispiel:

```
print(f"R^2 is: {lr.score(X_train, y_train):.3f}")
```

<!-- grade_id: answer3a, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Aufgabe 3b

Verwenden Sie exakt den Code von Aufgabe 3a. **Ersetzen Sie einzig** die Variable `Fruchtmasse` mit `Samen`. Lassen Sie den Code nochmal laufen.

<!-- grade_id: answer3b, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Was beobachten Sie? Erklären Sie.

<!-- grade_id: answer3b2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Aufgabe 3c

Verwenden Sie nun die Variablen `Fruchtmasse` und `Samen` gleichzeitig (_Multivariate Linear Regression_). **Ändern Sie was immer Sie für nötig halten (z.B. Modell-Parameter oder Pre-Processing)** um das Modell erfolgreich und mit möglichst wenig Gradient Descent Updates (steps) zu trainieren. Zeigen Sie mit einer geeigneten Grafik ob Batch Gradient Descent konvergiert hat (z.B. eine _Learning Curve_).

Folgende Parameter von `LinearRegression` sollen Sie verwenden und dürfen nicht geändert werden:
- `epsilon=1e-3`
- `alpha=0`
- `optimization_method="bgd"`
- `max_num_steps=x` wobei $x\le50000$ (damit Sie nicht ewig trainieren)


Versuchen Sie das Abbruchkriterium `Gradient < epsilon` zu erreichen.

<!-- grade_id: answer3c, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

Is es Ihnen gelungen das Abbruchkriterium zu erreichen? 

Falls Ja: Was war entscheidend?

Falls Nein: Warum nicht? Was hat Sie dem Ziel näher gebracht?

<!-- grade_id: answer3c2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Aufgabe 3d

Trainieren Sie nun zwei Modelle. Einmal mit `optimization_method="sgd"` und einmal mit `optimization_method="bgd"`. Sie können die restlichen Parameter selber wählen, sie sollen aber identisch sein für die beiden Modelle.

Zeichnen Sie für beide Modelle eine Learning-Curve. Wählen Sie eine Darstellung die einen einfachen Vergleich erlaubt.


<!-- grade_id: answer3d, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Wie unterscheiden sich die Learning-Curves? Warum?

<!-- grade_id: answer3d2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Aufgabe 3e

Nun untersuchen wir verschiedene Learning-Rates. Verwenden Sie dasselbe Setup wie in Aufgabe 3c. Mit folgenden Parameter:
- `epsilon=-1`
- `alpha=0`
- `optimization_method="bgd"`
- `max_num_steps=500`


Variieren Sie ausschliesslich den `learning_rate` Parameter in einem sinnvollen Bereich. Wählen Sie mind. 10 verschieden Werte. Trainieren Sie mehrere Modelle und vergleichen Sie die Konvergenz der Modelle für jeden Wert von `learning_rate`. Versuchen Sie eine möglichst hohe, funktionierende Learning-Rate zu finden. Erstellen Sie dazu Learning-Curves.

Achten Sie darauf, dass die Darstellung einen sinnvollen Vergleich erlaubt.

<!-- grade_id: answer3e, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
 
```

Interpretieren Sie die Learning-Curves. Was sehen Sie? Welche Learning-Rates funktionieren und sind praktikabel? Welche finden Sie optimal?

<!-- grade_id: answer3e2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Aufgabe 3f

In dieser Aufgabe untersuchen Sie den Effekt der Regularisierung $\lambda$ (Parameter `alpha`) auf das Modell. Dazu trainieren Sie mehrere Modelle mit verschiedenen Werten von `alpha` und vergleichen dann die Resultate. Folgende Bedingungen sollen Sie erfüllen:

- Verwenden Sie folgende Features: `"Samen", "Fruchtmasse", "Klongrösse", "Honigbiene"`

- Verwenden Sie `sklearn.preprocessing.PolynomialFeatures` um noch mehr Features zu generieren. Lesen Sie dazu die Dokumentation: [sklearn.preprocessing.PolynomialFeatures](https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.PolynomialFeatures.html)

- Wählen Sie die `alpha` über einen grosszügig gewählten Bereich. Alle anderen Parameter sollen für alle Modelle identisch sein.

- Vergleichen Sie für die verschiedenen Modelle die gefundenen Modell-Koeffizienten $\mathbf{w}$. Plotten Sie dazu `alpha` auf der X-Achse und den Wert von jedem Modell-Koeffizient auf der Y-Achse. Beispiel: [Link](https://scikit-learn.org/stable/auto_examples/linear_model/plot_ridge_path.html)


Tipp: Sie können `sklearn.preprocessing.StandardScaler` und `sklearn.pipeline.Pipeline`verwenden.


<!-- grade_id: answer3f, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Interpretieren Sie den Plot. Was sieht man und warum sieht es so aus?

<!-- grade_id: answer3f2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Aufgabe 4 (8 Punkte)

In dieser Aufgabe geht es darum ein best mögliches Modell zu finden.


### Aufgabe 4a

- Entwickeln Sie nun ein multiples lineares Regressionsmodell für die Zielgrösse `Ernte`.
- Sie dürfen durch Feature-Transformation beliebige weitere Attribute hinzufügen.
- Trainieren Sie das Modell mit dem Trainings-Datensatz. 
- Evaluieren Sie das beste Modell auf dem Trainings- und auf dem Testdatensatz.

Hinweis: In dieser Aufgabe geht es darum ein möglichst gutes Modell zu finden. Sie sollen demonstrieren, dass Sie in der Lage sind verschiedene Modell-Varianten miteinander zu vergleichen und die beste Variante auszuwählen. Ein ideales Instrument ist Kreuzvalidierung. Lesen Sie dazu die folgende Dokumentation: [cross_validation](https://scikit-learn.org/stable/modules/cross_validation.html). Verwenden Sie [sklearn.model_selection.GridSearchCV](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.GridSearchCV.html) um verschiedene Modelle miteinander zu vergleichen. 

Verwenden Sie geeignete Instrumente von sklearn, wie z.B.:
- [sklearn.preprocessing.StandardScaler](https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.StandardScaler.html)
- [sklearn.pipeline.Pipeline](https://scikit-learn.org/stable/modules/generated/sklearn.pipeline.Pipeline.html)
- [sklearn.compose](https://scikit-learn.org/stable/modules/classes.html#module-sklearn.compose)

<!-- grade_id: answer4a, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Aufgabe 4b

Interpretieren Sie das Ergebnis:

- Welche Modell-Varianten haben Sie verglichen? Welche ist die beste?
- Vergleichen und interpretieren Sie Scores auf Test- und Trainingsdatensatz. Wie schneidet das Modell zu den besten Modellen aus Aufgabe 3 ab?

<!-- grade_id: answer4b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Aufgabe 5 (2 Punkte)

In dieser Aufgabe geht es darum die Ergebnisse der einzelnen Teilaufgaben zu konsolidieren und ein Fazit zu ziehen.

### Aufgabe 5a

Beantworten Sie folgende Fragen:

- Wie schätzen Sie den Modellierungserfolg ein? Ist das Modell gut oder schlecht? Begründen Sie Ihre Antwort.

- Alex soll in einem nächsten Projekt herausfinden, welche Einflussgrössen wichtig sind. Was können Sie dazu bereits sagen?


<!-- grade_id: antwort5a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE